---
  title: "Voraussetzungen für statistische Tests"


---
## <<<New Exercise>>>

```yaml
type: TabExercise
key: f590749fec
lang: r
xp: 100
```


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

***

### Sub Heading

```yaml
type: NormalExercise
xp: 100
```

`@instructions`

`@solution`
```{r}

```

`@hint`

`@sct`
```{r}

```

***

### Sub Heading 2

```yaml
type: NormalExercise
xp: 100
```

`@instructions`

`@solution`
```{r}

```

`@hint`

`@sct`
```{r}

```
