---
  title: "Grafiken mit ggplot2"
  description: "Im ersten Teil hast du gelernt, wie du einfache Plots erstellen kannst. In diesem zweiten Teil wirst du lernen, wie du mit der Bibliothek ggplot2 Grafiken erstellst. Dies ist eine Bibliothek mit mehr Möglichkeiten in einfacher Art und Weise die Daten darzustellen."

---
## Paket laden

```yaml
type: NormalExercise 
xp: 100 
key: 38f9cedf15   
```

Wenn du eine Bibliothek verwenden möchtest, musst du sie erst laden. Das kannst du z.B. mit dem Befehl `library()`. In den Klammern hinter `library` steht der Name des Pakets. In dieser Lektion werden wir das für dich tun. Wenn du außerhalb dieser Lektion R-Code mit den hier verwendeten Befehlen ausführen willst, solltest du die entsprechenden Bibliotheken laden! Wenn du dir nicht sicher bist, welches Paket du laden musst, kannst du in der Hilfe  nach dem Befehl, den du benutzen möchtest, suchen und siehst, entweder in geschweiften Klammern hinter den Befehl auf einer Hilfeseite oder falls mehrere Befehle gefunden werden, die zu deiner Suchanfrage passen, vor den doppelten Doppelpunkten den Paketnamen.

`@instructions`
Lade das Paket `ggplot2`.

`@hint`
Gib `?library` ein.

`@pre_exercise_code`

```{r}

```

`@sample_code`

```{r}

```

`@solution`

```{r}

```

`@sct`

```{r}

```





